---
title       : Insert the chapter title here
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf


--- type:VideoExercise lang:r xp:50 skills:1 key:c887bfcaa6
## Test

*** =projector_key
243f10cc189e54f8e8c51edd8ee468f0

--- type:NormalExercise lang:r xp:100 skills:1 key:0d449f94b2
## Test


*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
